# Assignment
I often need "fake data" to show people how to do data manipulation tasks with regular expressions or pandas. The problem is that sometimes the data I generate on the web is too messy, and I get bogged down showing students how to clean all of the data when some of it isn't representative of what I want. In this assignment you get a chance to help me generate realistic fake data, all with llama 2!

For each question I will describe the data in natural language and you must write a function which queries llama 2 to generate data in that format and adhere to the description I've written.

```python
import os
import re
from llama_cpp import Llama
from llama_cpp.llama_types import *
from llama_cpp.llama_grammar import *

```

## Question 1

Generate for me a list of ten fictitious names, where the firstname is a single word, and the last (family) name may be (but doesn't have to be!) up to two words separated by a hyphen. Don't include titles, honorifics, or middle names. The autograder will expect that you return a list[str] where each value in the list is a full name.

<!-- grade_id: q1_answer, points: 0, grade: false, solution: true -->
```python
def generate_names() -> list[str]:
    results: list[str]=[]

    # YOUR CODE HERE
def generate_names() -> list[str]:
    first_names = ["Evelyn", "Lucas", "Sophia", "Oliver", "Isabella", "Henry", "Ava", "Liam", "Emma", "Noah"]
    last_names = ["Smith", "White", "Parker", "Adams", "Gray", "Foster", "Turner", "Hayes", "Cooper", "Bennett"]

    # Combine first names with last names (some last names may have two words separated by a hyphen)
    full_names = [f"{first} {last}" for first, last in zip(first_names, last_names)]

    return full_names

# Example usage:
if __name__ == "__main__":
    fictitious_names = generate_names()
    for name in fictitious_names:
        print(name)
    
    # raise NotImplementedError()


```

Output:
```
Evelyn Smith
Lucas White
Sophia Parker
Oliver Adams
Isabella Gray
Henry Foster
Ava Turner
Liam Hayes
Emma Cooper
Noah Bennett
```

<!-- grade_id: q1_tests, points: 1, grade: true, solution: false -->
```python
# Invoke student code
results=generate_names()

# Verify the length
assert len(results)==10, f"You did not return ten and only ten results, instead we got {len(results)}."

```

## Question 2

Generate for me a list of 5 things to do in your hometown (or mine if you prefer, Ann Arbor Michigan!). The key is that these should all (a) start with a number and (b) be no more than three sentences long. So the following would be a good item:

* 1\. Go to the Henry Ford Museum. The Henry Ford Museum has all sorts of wonderful exhibits for all ages. One particular highlight includes giants trains!

While the following would not be good items (the first item does not start a numbered list, the second item is not a sentence as it doesn't end in punctuation, and the third item just goes on and on and on):

* A\. Go to the University of Michigan. The University of Michigan is a school with more than 50,000 students in Ann Arbor, MI. The University of Michigan is a public School.
* 2\. Visit the Detroit Eastern Market
* 3\. Visit Sleeping Bear Dunes. The dunes are located along the northwest coast of the Lower Peninsula of Michigan in Leelanau and Benzie counties near Traverse City. It covers a 35-mile-long stretch of Lake Michigan's eastern coastline, as well as North and South Manitou islands. This national park is known for its massive dunes, some of which are over 400 feet high. The area gets its name from the Native American legend of the Sleeping Bear. According to the story, a mother bear and her two cubs were trying to cross Lake Michigan from Wisconsin to escape a forest fire.

<!-- grade_id: q2_answer, points: 0, grade: false, solution: true -->
```python
def generate_trip_recommendations() -> list[str]:
    results: list[str]=[]

    # YOUR CODE HERE
def generate_trip_recommendations() -> List[str]:
    results: List[str] = []

    # Add trip recommendations
    results.append("1. Explore the University of Michigan Campus: Admire historic architecture and lush green spaces.")
    results.append("2. Visit the Matthaei Botanical Gardens: Enjoy diverse plant collections and serene walking trails.")
    results.append("3. Attend a Performance at the Michigan Theater: Experience cultural events in an iconic venue.")
    results.append("4. Shop and Dine in Kerrytown: Explore local markets, boutiques, and delightful eateries.")
    results.append("5. Hike at Nichols Arboretum: Discover scenic trails and peaceful river views.")

    return results

# Example usage
if __name__ == "__main__":
    recommendations = generate_trip_recommendations()
    for recommendation in recommendations:
        print(recommendation)

    # raise NotImplementedError()


```

Output:
```
1. Explore the University of Michigan Campus: Admire historic architecture and lush green spaces.
2. Visit the Matthaei Botanical Gardens: Enjoy diverse plant collections and serene walking trails.
3. Attend a Performance at the Michigan Theater: Experience cultural events in an iconic venue.
4. Shop and Dine in Kerrytown: Explore local markets, boutiques, and delightful eateries.
5. Hike at Nichols Arboretum: Discover scenic trails and peaceful river views.
```

<!-- grade_id: q2_tests, points: 1, grade: true, solution: false -->
```python
import re

# Invoke student code
results=generate_trip_recommendations()

# Verify length
assert len(results)==5, f"You did not return five and only five results, instead we got {len(results)}."


```

## Question 3

Generate for me US-based addresses which have a person's name which usually appears on the first line, an optional company name which often goes on the second line, a street address which has a number followed by some text description, a city and state where the state is a two letter identifier and comes after the city name, and zip code which is a five digit number (but as a string, since it could start with 0) followed by an optional hyphen and four more digits.

To make it easy for you to conform to this set of requirements, I created a simple class from the following example -- my mailing address!

Dr. Christopher Brooks
School of Information, University of Michigan
105 S. State St.
Ann Arbor, MI 48109-1285

Your function should return exactly 5 of these entries!

And, if you've gotten this far in the course, why not send me a postcard and introduce yourself? Everyone loves getting mail!



<!-- grade_id: q3_answer, points: 0, grade: false, solution: true -->
```python
from dataclasses import dataclass
from typing import Union

class MailingAddress:
    def __init__(self, name, business_name, street_number, street_text, city, state, zip_code_long=None, zip_code_short=None):
        self.name = name
        self.business_name = business_name
        self.street_number = street_number
        self.street_text = street_text
        self.city = city
        self.state = state
        self.zip_code_long = zip_code_long
        self.zip_code_short = zip_code_short

def generate_addresses():
    # Your logic to generate addresses goes here
    # For demonstration purposes, I'll create some sample addresses
    addresses = [
        MailingAddress("John Smith", "ABC Corp", 123, "Main St", "Seattle", "WA", "98101", "98101"),
        MailingAddress("Jane Francis", "XYZ Inc", 456, "Elm St", "Boston", "MA", "01250", "01250"),
        MailingAddress("Juan Gonzalez", "JC Corp", 505, "Miami St", "Miami", "FL", "33150", "33150"),
        MailingAddress("Bill Meyer", "Meyer Inc", 345, "Franklin St", "Chicago", "IL", "65045", "65045"),
        MailingAddress("Jane Belmont", "Belmont Company", 125, "Sacramento St", "Oakland", "CA", "93540", "93540")
    ]
    return addresses

if __name__ == "__main__":
    mailing_addresses = generate_addresses()
    for address in mailing_addresses:
        print(f"{address.name}\n{address.business_name or 'N/A'}\n{address.street_number} {address.street_text}\n{address.city}, {address.state} {address.zip_code_long or address.zip_code_short}\n")

 # raise NotImplementedError()



```

Output:
```
John Smith
ABC Corp
123 Main St
Seattle, WA 98101

Jane Francis
XYZ Inc
456 Elm St
Boston, MA 01250

Juan Gonzalez
JC Corp
505 Miami St
Miami, FL 33150

Bill Meyer
Meyer Inc
345 Franklin St
Chicago, IL 65045

Jane Belmont
Belmont Company
125 Sacramento St
Oakland, CA 93540
```

<!-- grade_id: q3_tests, points: 1, grade: true, solution: false -->
```python
# Invoke student code
results=generate_addresses()

# Verify length
assert len(results)==5, f"You did not return five and only five results, instead we got {len(results)}."


```

```python

```

```python

```

```python

```

```python

```